# COGS 108 - Assignment 1: git & Python

This assignment covers git, GitHub, and Python Basics. 

Feel free to check out the relevant tutorials on https://github.com/COGS108/Tutorials and available [Resources](https://github.com/COGS108/Resources) for guidance for these topics if you're less familiar.

This assignment is out of 8 points, worth 8% of your grade.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.** (Note that you *can* add additional cells, if you want to test things out.)

## How to complete assignments

Whenever you see:

```python
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (meaning: delete) these lines of code with some code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

### How to Score Points

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question, and there may be some additional hidden `assert` statements to test your code.

Reading assert errors if and when they fail is a great way to debug your code and find out what went wrong.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

All outside packages/modules that will be used will be specified. You may not use other libraries in the assignments.

Finally, note that questions have points as specified in the detailed instructions. 

## Introduction

The purpose of this assignment is to make sure you have the tools you'll need to use
for COGS 108. Notably, we'll be using Python, Jupyter notebooks, and git/GitHub. Since we're using [datahub](http://datahub.ucsd.edu), you won't *need* a local version of Jupyter or Python on your computer. So, for this assignment, we'll focus on getting you set up on GitHub.

Most assignments will be completed completely within the Jupyter Notebook and submitted on datahub; however, for this first assignment, we want you to be comfortable in GitHub. Tasks to complete Part 1 below will require work outside of this notebook.


## Part 1: git & GitHub (2.5 points)

If you're new to git and GitHub, documentation [here](https://help.github.com/en/github/getting-started-with-github), [here](https://guides.github.com/activities/hello-world/), and [here](https://github.com/jlord/git-it-electron#what-to-install) will likely be helpful. 

Below, we will have you accomplish three common GitHub tasks. Well explain them briefly here and _why_ we're having you do them. Step-by-step instructions are then included below.

* Cloning repositories from GitHub - you'll do this any time you want to get a repo from GitHub to your local machine.
* Creating a new repository on your GitHub - here, you'll create the repo, add files to the repo, and commit these changes
* Fork a repo & submit a PR - any time you want to work with someone else's repo to which you don't have access, you'll use this procedure. 

#### Step-by-Step Instructions

1. If you do not already have one, **create a GitHub account**
  - Follow instructions here: https://github.com/
2. **Make a public repo** on your GitHub called 'COGS108_Repo'
  - Make sure this repository is public
  - Make sure you name the repo exactly as you see it above
  - Within `COGS108_Repo`, create two files: 'README.md' and '.gitignore'. 
      - Add these to 'COGS108_Repo'.
      - Reminder: the shell command `touch` creates an empty file.
      - These can contain anything you want (or be empty).
      - Stage, commit, and push these changes to your GitHub.
3. **Make a Pull Request**
  - Fork the COGS108 repo called 'MyFirstPullRequest' to your GitHub account.
      - Go to the COGS108/MyFirstPullRequest on GitHub. 
          - On the top right of the page, click on the 'Fork' button and click through to Fork the repository to your GitHub account
  - Clone that fork onto your computer.
  - You are going to add a text file with a unique identifier to your repo
      - The unique identifier should be: the first letter of your first name (capitalized) + the first letter of your last name (capitalized) + the last 2 digits of your Student ID (PID). 
      - For example, John Smith - A12345678 should submit 'JS78.txt'. 
      - The contents of this file do not matter (it can be empty). 
      - Add this file to your cloned repository (the folder on your computer).
      - Note: A 'txt' file is a plain text file. It should have the extension '.txt'. You may need to specify, in your preferred text editor, to save this filetype.
  - Stage & Commit these files.
  - Push these updates to the cloned repo, on your GitHub account.
  - Go to GitHub.com, to '/YourAcct/MyFirstPullRequest' where you should now see your txt file. 
  - Create a Pull Request (PR).
    - Click on 'New Pull Request'
    - Click on 'Create Pull Request' (in green)
    - Replace the commit message with the same identifier (ex - 'JS78')
    - Click 'Create Pull Request'.
4. *Optional*: **Clone the class repositories** (We won't be grading this...but if you want to have access to the course materials on your computer, it may be a good idea.)
  - Go to the class organization: https://github.com/COGS108
  - Clone 'Lectures-Ellis' and 'Projects'
- You're done!
  
#### General Notes

Make sure all the specified repos and files that you make on GitHub are *public*, and that you type your GitHub username in properly below. Automated procedures will be used to check that these repos exist, that the Pull Request was received, and that the submitted Jupyter notebook runs as expected. 


### Q0: <span style="color: red;">PID & GitHub Username </span> (2.5 points)

Once you've completed the above instructions, we'll need you to provide us with your PID and GitHub username so that we can check completion and accuracy of the tasks assigned above.

In the cell below, assign your UCSD PID as a string to the variable `PID` and your GitHub username as a string to the variable `github_username`.

For example, the cell below may look something like this (But will include *your* PID and username):

```python
PID = 'A1234567'
github_username = 'ShanEllis'
```

*Note*: Make sure you provide your GitHub username, **not the email address** you used to create the account, follow instructions, and avoid typos. You will not get credit for this section if typos are made.

<!-- grade_id: cell-784114344a572182, points: 0, grade: false, solution: true -->
```python
# your code here
PID = "A16740028"
github_username = "kisenera"
```

```python
# if you navigate to this link
# it should open your GitHub page
# if it doesn't, fix your username above
'www.github.com/' + github_username
```

Output:
```
'www.github.com/kisenera'
```

<!-- grade_id: cell-216ded1c5ab2c280, points: 0, grade: true, solution: false -->
```python
assert PID
assert github_username
```

## Part 2: Python (5.5 points)

This part of the assignment is focused on some practice with Python, and with practicing working with the format of the assignments. 

This class assumes some prior knowledge of Python. In the following questions, you will need to work with basic (standard library) data types (floats, lists, dictionaries, etc.) and control flow (conditionals, loops, functions, etc). If the questions in this section are totally unfamiliar to you, you may need to revisit some practice materials to catch up with some of the programming. 

Through these questions, we will also prompt you to use a couple slightly more advanced standard library functions (for example, `enumerate` and `zip`), that may be new to you.

Each question should be answerable with a relatively small number of lines of code, up to about 5-7 lines at most. 

If you are having any trouble, remember to visit the course tutorials (https://github.com/COGS108/Tutorials) and/or resources (https://github.com/COGS108/Resources). Assignment questions often follow the structure of examples provided in the tutorials, and a large number of relevant links and external materials are also indexed in the tutorials. 

```python
# PRINTING VARIABLES
# A reminder that you can (and should) print and check variables as you go.
#  This allows you to check what values they hold, and debug if anything unexpected happens.

# Define a variable
math_result = 2 * 4

# Print out the value(s) of a variable.
print(math_result)
```

Output:
```
8
```

**Q1: Defining Variables (0.25 points)**

Create the following variables:
- a variable called `my_int` that stores the integer 2
- a variable called `my_float` that stores the float 2.29
- a variable called `my_string` that stores the string 'COGS'
- a variable called `my_bool` that stores the boolean `False`

<!-- grade_id: cell-ba64d778a1562c6e, points: 0, grade: false, solution: true -->
```python
# your code here
my_int = 2
my_float = 2.29
my_string = 'COGS'
my_bool = False
```

<!-- grade_id: cell-6b812a82cf83cb02, points: 0.25, grade: true, solution: false -->
```python
# Tests that check that each variable is of the right type
assert isinstance(my_int, int)
assert isinstance(my_float, float)
assert isinstance(my_string, str)
assert isinstance(my_bool, bool)

```

**Q2: Defining Variables: Lists & Tuples (0.25 points)**

Define a list called `var_a`, that contains individual letters a-f (inclusively).

Define a tuple called `var_b`, that contains the numbers 0-5 (inclusively).

<!-- grade_id: cell-4db8c933a259ed1b, points: 0, grade: false, solution: true -->
```python
# your code here
var_a = ['a', 'b', 'c', 'd', 'e', 'f']
var_b = (0, 1, 2, 3, 4, 5)
```

<!-- grade_id: cell-88582fbb27969b63, points: 0.25, grade: true, solution: false -->
```python
# These tests check the variables are defined
assert var_a
assert var_b

# These tests check that the variables are the right data types
assert isinstance(var_a, list)
assert isinstance(var_b, tuple)

# hidden tests to make sure you have the right values in the variables
### BEGIN HIDDEN TEST
# These tests check the variables have the right values
assert var_a == ['a', 'b', 'c', 'd', 'e', 'f']
assert var_b == (0, 1, 2, 3, 4, 5)
### END HIDDEN TEST
```

**Q3: Defining Variables: Dictionaries (0.5 points)**

Create a Python dict called `dictionary` where the keys are the elements in `var_a` and the values are the corresponding elements in `var_b`. (Note: Use `var_a` and `var_b` you created above.)

Hints:
1. The `zip` function may be useful.
2. You might also make use of a Python dictionary comprehension. (It's OK if you use a different approach.)

<!-- grade_id: cell-a0c9ee8ea65481d3, points: 0, grade: false, solution: true -->
```python
# your code here
dictionary = dict(zip(var_a, var_b))
```

<!-- grade_id: cell-611de86b26434b5e, points: 0.5, grade: true, solution: false -->
```python
assert dictionary
assert isinstance(dictionary, dict)

# There are also some hidden tests that check the contents of dictionary
```

**Q4: Value Comparisons (0.25 points)**

Store values in the variables `comp_val_1`, `comp_val_2`,`comp_val_3`, and `comp_val_4` such that the assert cell will pass silently (meaning the tests will pass and not produce an error). 

<!-- grade_id: cell-2117f36254285222, points: 0, grade: false, solution: true -->
```python
# your code here
comp_val_1 = 2
comp_val_2 = 8
comp_val_3 = 1
comp_val_4 = 1
```

<!-- grade_id: cell-68658c3ed892c21b, points: 0.25, grade: true, solution: false -->
```python
assert comp_val_1 < comp_val_2
assert comp_val_3 == comp_val_4
assert comp_val_4 <= comp_val_1
```

**Q5: Control Flow (0.5 points)**

Loop through the provided list `my_list`. For each element, check if it is an odd number. If the element is an odd number, append the INDEX of that element to the list `inds`.

Note that you are adding the index to `inds`, **not the element itself**. (Reminder: Python uses zero-based indexing, so the index of the first element in the list is 0.)
   
Hints: 
1. To check if a number is even, you can use the modulo `%` operator.
2. To loop through an iterable, keeping track of the index, you can (but are not required to) use the `enumerate` function.

```python
# These variables are provided to you.
my_list = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]
inds = []
```

<!-- grade_id: cell-5c803671612c731a, points: 0, grade: false, solution: true -->
```python
# your code here
for i, e in enumerate(my_list):
    if i % 2 != 0:
        inds.append(e)
```

<!-- grade_id: cell-eec717db577f9cce, points: 0.25, grade: true, solution: false -->
```python
assert isinstance(inds, list)
```

<!-- grade_id: cell-90cd314531614439, points: 0.25, grade: true, solution: false -->
```python
assert inds

# There is also a hidden test that checks the contents of 'inds'
```

**Q6: Indexing (0.5 points)**

Using the four lists provided in the cell below, complete the following *using indexing/slicing*:
- Use forward indexing to store the second value in `list_1` to `index_1`
- Use negative indexing to store the second value from the end in `list_2` to `index_2`
- Use slicing to store the third and fourth values (inclusive) of `list_3` to `index_3`
- Use slicing to store the last two values of `list_4` to `index_4` 

Note: Slices should store information in the same order as the original list.

```python
# provided for you
list_1 = [10, 20, 30, 40]
list_2 = [13, 15, 17, 19, 21, 23]
list_3 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
list_4 = [21, 9, 98, 289, 938]
```

<!-- grade_id: cell-4094a1f93e0d8bd7, points: 0, grade: false, solution: true -->
```python
# your code here
index_1 = list_1[1]
index_2 = list_2[-2]
index_3 = list_3[2:4]
index_4 = list_4[-2:]
```

<!-- grade_id: cell-4bf3bc61f6f0e1b8, points: 0.1, grade: true, solution: false -->
```python
assert isinstance(index_1, int)
```

<!-- grade_id: cell-9d8e61c652862c50, points: 0.1, grade: true, solution: false -->
```python
assert isinstance(index_2, int)
```

<!-- grade_id: cell-f6562819b4f97063, points: 0.15, grade: true, solution: false -->
```python
assert isinstance(index_3, list)
```

<!-- grade_id: cell-a93b32d8d0cf0352, points: 0.15, grade: true, solution: false -->
```python
assert isinstance(index_4, list)
```

**Q7: Looping through Dictionaries (0.75 points)**

Imagine, for example, we have a `students` dictionary that indicates whether a student has completed an assignment and we want to track the students that have not yet completed the assignment.

Using the `students` dictionary provided below, write a for loop that loops across the dictionary and collects all subject numbers (ex. 'S2') where the dictionary value is `False`. Store the subject numbers in a new list named `incomplete`.

The code should programmatically determine which keys store the value `False`, storing those keys in a list.

```python
# This dictionary provided for you
students = {
    'S0' : True,
    'S1' : True,
    'S2' : False,
    'S3' : True,
    'S4' : False,
    'S5' : True
}
```

<!-- grade_id: cell-293277807c63baaf, points: 0, grade: false, solution: true -->
```python
# your code here
incomplete = []
for k, v in students.items():
    if v == False:
        incomplete.append(k)
```

<!-- grade_id: cell-2a806b0b0c85649a, points: 0.75, grade: true, solution: false -->
```python
assert isinstance(incomplete, list)
assert len(incomplete) == 2

```

**Q8: Functions I (0.75 points)**

Write a function `return_even` that will take a list as its input and return a list of all the even *values* as its output. 

Note that this differs from what you did above in two ways: (1) it returns the values, not the index, and (2) it's looking for even values, not odd.

For example:

```python
>>> return_even([1, 2, 3, 4])
[2, 4]
>>> return_even([1, 3, 5])
[]
```

<!-- grade_id: cell-f5cd038c9d5c81a2, points: 0, grade: false, solution: true -->
```python
# your code here
def return_even(l):
    ans = []
    for i in l:
        if i % 2 == 0:
            ans.append(i)
    return ans
```

<!-- grade_id: cell-72092ce37bd26a1b, points: 0.75, grade: true, solution: false -->
```python
assert return_even([1, 3, 5]) == []
assert return_even([1, 2, 3, 4]) == [2, 4]
```

**Q9: Functions II (0.75 points)**

Write a function `squared_diff` that takes two number inputs and returns the squared difference of the two numbers i.e., $(a - b)^2$. For example:

```python
>>> squared_diff(2, 4)
4
>>> squared_diff(10, 1)
81
```

<!-- grade_id: add2, points: 0, grade: false, solution: true -->
```python
# your code here
def squared_diff(a, b):
    c = a - b
    return c ** 2
```

<!-- grade_id: add2-tests, points: 0.75, grade: true, solution: false -->
```python
assert squared_diff(2, 4) == 4
assert squared_diff(10, 1) == 81

# There are also hidden tests that check more cases of squared_diff
```

**Q10: Putting it all together (1 point)**

Here, we'll update the values in `dictionary`, storing the output in a dictionary called `other_dictionary`. This new dictionary will have the same keys, but some values will be updated.

The values in `other_dictionary` should be updated, such that if the value in the original `dictionary` is...
* odd: update the the value stored in the dictionary to store the squared difference of the original value and '10'. Remember, you created a function to do this above.)
* even: store the original value (from `dictionary`). 

Hints: 
1. to loop through key-value pairs in a dictionary, check out the `.items` method
2. You created a `squared_diff` function above.

```python
# run this cell to remind yourself what is stored in `dictionary`
dictionary.items()
```

Output:
```
dict_items([('a', 0), ('b', 1), ('c', 2), ('d', 3), ('e', 4), ('f', 5)])
```

<!-- grade_id: cell-50b7166b6554cfff, points: 0, grade: false, solution: true -->
```python
# your code here
other_dictionary = dictionary
for k, v in other_dictionary.items():
    if v % 2 != 0:
        other_dictionary[k] = squared_diff(v, 10)
```

<!-- grade_id: cell-efa2dd1099f123b0, points: 0.25, grade: true, solution: false -->
```python
assert isinstance(other_dictionary, dict)
```

<!-- grade_id: cell-234cdbef12134b0d, points: 0.75, grade: true, solution: false -->
```python
assert other_dictionary

# There are hidden tests that check the contents of 'other_dictionary'.
```

## The End!

This is the end of the first assignment!

The goal here was to check your basic Python understanding. From here, assignments will be more data science centric, using lots of `pandas` and working with data!

Have a look back over your answers, and also make sure to `Restart & Run All` from the kernel menu to double check that everything is working properly. You can also use the 'Validate' button above, which runs your notebook from top to bottom and checks to ensure all `assert` statements pass silently. When you are ready, submit on datahub!

```python

```